The BTB-ZF gene Bm-mamo regulates pigmentation in silkworm caterpillars

Author Response

The following is the authors’ response to the previous reviews.

eLife assessment

This important study identifies the gene mamo as a new regulator of pigmentation in the silkworm Bombyx mori, a function that was previously unsuspected based on extensive work on Drosophila where the mamo gene is involved in gamete production. The evidence supporting the role of Bm-nano in pigmentation is convincing, including high-resolution linkage mapping of two mutant strains, expression profiling, and reproduction of the mutant phenotypes with state-of-the-art RNAi and CRISPR knock-out assays. While the discussion about genetic changes being guided or accelerated by the environment is extremely speculative and has little relevance for the findings presented, the work will be of interest to evolutionary biologists and geneticists studying color patterns and evolution of gene networks.

Response: Thank you very much for your careful work. In the revised version, we conducted a comparative genomic analysis of the upstream regions of the Bm-mamo gene in 51 wild silkworms and 171 domesticated local silkworms. The analysis of nucleotide diversity (pi) and the fixation index (FSTs) of the Bm-mamo genome sequences in the wild and domesticated silkworm populations were also performed. The results showed that the Bm-mamo genome sequence of local silkworms was relatively conserved, while the upstream sequence of wild silkworms exhibited high nucleotide diversity. This finding suggested a high degree of variability in the regulatory region of the Bm-mamo gene, in wild strains. Additionally, the sequence in this region may have been fixed by domestication selection. We have optimized the description in the discussion section.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

This papers performs fine-mapping of the silkworm mutants bd and its fertile allelic version, bdf, narrowing down the causal intervals to a small interval of a handful of genes. In this region, the gene orthologous to mamo is impaired by a large indel, and its function is later confirmed using expression profiling, RNAi, and CRISPR KO. All these experiments are convincingly showing that mamo is necessary for the suppression of melanic pigmentation in the silkworm larval integument.

The authors also use in silico and in vitro assays to probe the potential effector genes that mamo may regulate.

Strengths:

The genotype-to-phenotype workflow, combining forward (mapping) and reverse genetics (RNAi and CRISPR loss-of-function assays) linking mamo to pigmentation are extremely convincing.

This revision is a much improved manuscript and I command the authors for many of their edits.

Response: Thank you very much for your careful work. With the help of reviewers and editors, we have revised the manuscript to improve its readability.

I find the last part of the discussion, starting at "It is generally believed that changes in gene expression patterns are the result of the evolution of CREs", to be confusing.

In this section, I believe the authors sequentially:

emphasize the role of CRE in morphological evolution (I agree)

emphasize that TF, and in particular their own CRE, are themselves important mutational targets of evolution (I agree, but the phrasing need to insist the authors are here talking about the CRE found at the TF locus, not the CRE bound by the TF).

use the stickleback Pel enhancer as an example, which I think is a good case study, but the authors also then make an argument about DNA fragility sites, which is hard to connect with the present study.

then continue on "DNA fragility" using the peppered moth and butterfly cortex locus. There is no evidence of DNA fragility at these loci, so the connection does not work. "The cortex gene locus is frequently mutated in Lepidoptera", the authors say. But a more accurate picture would be that the cortex locus is repeatedly involved in the generation of color pattern variants. Unlike for Pel fragile enhancer, we don't know if the causal mutations at this locus are repeatedly the same, and the haplotypes that have been described could be collateral rather than causal. Overall, it is important to clarify the idea that mutation bias is a possible factor explaining "genetic hotspots of evolution" (or genetic parallelism sensu 10.1038/nrg3483), but it is also possible that many genetic hotspots are repeated mutational targets because of their "optimal pleiotropy" (e.g. hub position in GRNs, such as mamo might be), or because of particularly modular CRE region that allow fine-tuning. Thus, I find the "fragility" argument misleading here. In fact the finding that "bd" and "bdf" alleles are different in nature is against the idea of a fragility bias (unless the authors can show increased mutation rates at this locus in a wild silkmoth species?). These alleles are also artificially-selected ie. they increased in frequency by breeding rather than natural selection in the wild, so while interesting for our understand of the genotype-phenotype map, they are not necessarily representative of the mutations that may underlie evolution in the wild.

Response: Thank you very much for your careful work. DNA fragility is an interesting topic, but some explanations for DNA fragility are confusing. One study measured the rate of DNA double-strand breaks (DSBs) in yeast artificial chromosomes (YACs), which are chromosomes containing marine Pel that broke ~25 to 50 times more frequently than did the control. These authors believe that the increase in the mutation rate is caused by DNA sequence characteristics, particularly TG-dinucleotide repeats. Moreover, they found that adding a replication origin on the opposite side of Pel did not cause the fungus to switch fragile, making the forward sequence stable and the reverse complement fragile. Thus, Pel fragility is also dependent on the direction of DNA replication. In summary, they suggested that the special DNA sequence is the cause of DNA fragility. In addition, the sequence features associated with DNA fragility in the Pel region are also found in thousands of other positions in the stickleback and human genomes (Xie KT et al, 2019, science).

In yeast artificial chromosomes (YACs), the characteristics of DNA sequences, such as TG-dinucleotide repeat sequences, may be important reasons for DNA fragility, and these breaks occur during DNA replication. However, the inserted sequence of YAC often undergoes deletion or recombination during cultivation and passage. In addition, yeast is a single-celled organism. Therefore, the results in yeast cannot represent the situation in multicellular organisms. If multicellular organisms are like this, there are several issues as follows:

(1) The DNA replication process occurs separately in different multicellular organisms. Because DNA breakage and repair are independent, they can lead to the presence of different alleles in different cells. This can potentially lead to the occurrence of extensive chimeric organisms. However, we have not found such a situation in the genome sequencing of many multicellular organisms.

(2) If the DNA sequence, TG-dinucleotide repeats, is the determining factor, the mutations near the sequence lose their strong correlation with environmental changes. The researchers conducted yeast artificial chromosome experiments in the same environment and found that the frequency of DNA breaks containing TG dinucleotide repeat sequences was 25 to 50 times greater than that of the control group. This means that, whether in the marine population or the lake population, this part of the sticklebacks’ genome has undergone frequent mutations. However, according to related research, populations of lake sticklebacks, rather than marine populations, often exhibit a decrease in the pelvic phenotype.

(3) Researchers have found thousands of loci in the genome of sticklebacks and humans that contain such sequences (TG-dinucleotide repeats). This means that thousands of sites undergo frequent mutations during DNA replication. Unless these sites do not possess functionality, they will have some impact on the organism, even causing damage. Even if they are not functional sequences, these sequences will gradually be discarded or replaced during frequent mutations rather than being present in large quantities in the genome.

Therefore, the study of DNA fragility in yeast cannot explain the situation in multicellular organisms.

As you noted, we want to express that the frequent variation in the cortex gene should be regulated by targeted regulation involving the GRN in Lepidoptera. In addition, studies on specific epigenetic modifications discovered through the referenced fragile DNA sites suggest that DNA fragility is not determined by the DNA sequence (Ji F, 2020, Cell Res) but rather by other factors, such as epigenetic factors. The sequence features discovered at fragile DNA sites are traces of frequent mutations, not causes.

In this revision, we analyzed the nucleotide diversity of the mamo genome in 51 wild and 171 domestic silkworms. We found high nucleic acid diversity from the third exon to the upstream region of this gene in wild silkworms. We randomly selected 12 wild silkworms and 12 domestic silkworms and compared their upstream sequences to approximately 1 kb. In wild silkworms, there is significant diversity in their upstream sequences. In domestic silkworms, the sequences are highly conserved, but in some silkworms, a long interspersed nuclear element (LINE) is inserted. This finding suggested that there is frequent variation in the sequence of this region in wild silkworms, while fixation occurs in domesticated silkworms. These genomic data are sourced from the pangenome of silkworms (Tong X, 2022, Nat Commun.). In the pangenomic research, 1078 strains (205 local strains, 194 improved strains, 632 mutant strains, and 47 wild silkworms), which included 545 third-generation sequencing genomes, were obtained. An online website was built to utilize these data (http://silkmeta.org.cn/). We warmly welcome you to use these data.

In summary, for clearer expression, we have rewritten this section.

Xie KT, Wang G, Thompson AC, Wucherpfennig JI, Reimchen TE, MacColl ADC, Schluter D, Bell MA, Vasquez KM, Kingsley DM. DNA fragility in the parallel evolution of pelvic reduction in stickleback fish. Science. 2019 Jan 4;363(6422):81-84. doi: 10.1126/science.aan1425.

Ji F, Liao H, Pan S, Ouyang L, Jia F, Fu Z, Zhang F, Geng X, Wang X, Li T, Liu S, Syeda MZ, Chen H, Li W, Chen Z, Shen H, Ying S. Genome-wide high-resolution mapping of mitotic DNA synthesis sites and common fragile sites by direct sequencing. Cell Res. 2020 Nov;30(11):1009-1023. doi: 10.1038/s41422-020-0357-y.

Tong X, Han MJ, Lu K, Tai S, Liang S, Liu Y, Hu H, Shen J, Long A, Zhan C, Ding X, Liu S, Gao Q, Zhang B, Zhou L, Tan D, Yuan Y, Guo N, Li YH, Wu Z, Liu L, Li C, Lu Y, Gai T, Zhang Y, Yang R, Qian H, Liu Y, Luo J, Zheng L, Lou J, Peng Y, Zuo W, Song J, He S, Wu S, Zou Y, Zhou L, Cheng L, Tang Y, Cheng G, Yuan L, He W, Xu J, Fu T, Xiao Y, Lei T, Xu A, Yin Y, Wang J, Monteiro A, Westhof E, Lu C, Tian Z, Wang W, Xiang Z, Dai F. High-resolution silkworm pan-genome provides genetic insights into artificial selection and ecological adaptation. Nat Commun. 2022 Sep 24;13(1):5619. doi: 10.1038/s41467-022-33366-x.

Lu K, Pan Y, Shen J, Yang L, Zhan C, Liang S, Tai S, Wan L, Li T, Cheng T, Ma B, Pan G, He N, Lu C, Westhof E, Xiang Z, Han MJ, Tong X, Dai F. SilkMeta: a comprehensive platform for sharing and exploiting pan-genomic and multi-omic silkworm data. Nucleic Acids Res. 2024 Jan 5;52(D1):D1024-D1032. doi: 10.1093/nar/gkad956.

Curiously, the last paragraph ("Some research suggests that common fragile sites...") elaborate on the idea that some sites of the genome are prone to mutation. The connection with mamo and the current article are extremely thin. There is here an attempt to connect meiotic and mitotic breaks to Bm-mamo, but this is confusing: it seems to propose Bm-mamo as a recruiter of epigenetic modulators that may drive higher mutation rates elsewhere. Not only I am not convinced by this argument without actual data, but this would not explain how the mutations at the Bm-mamo itself evolved.

Response: Thank you very much for your careful work. This section mainly illustrates that DNA fragility is not determined by sequence but is regulated by other factors in animals. In fruit flies, they found that mamo is an important candidate gene for recombination hotspot setting in meiosis. First, we evaluated PRDM9, which plays an important role in setting recombination hotspots during meiosis. Our purpose in mentioning this information is to illustrate that chromosome recombination is a process of programmed double strand breaks and to answer another reviewer's question about programmed events in the genome. In summary, we suggest that some variations in DNA sequences are procedural results. We have optimized the description of this section in this version.

On a more positive note, I find it fascinating that the authors identified a TF that clearly articulates or orchestrate larval pattern development, and that when it is deleted, can generate healthy individuals. In other words, while it is a TF with many targets, it is not too pleiotropic. This idea, that the genetically causal modulators of developmental evolution are regulatory genes, has been described elsewhere (e.g. Fig 4c in 10.1038/s41576-020-0234-z, and associated refs). To me, the beautiful findings about Bm-mamo make sense in the general, existing framework that developmental processes and regulatory networks "shape" the evolutionary potential and trajectories of organisms. There is a degree of "programmability" in the genomes, because some loci are particularly prone to modulate a given type of trait. Here, Bm-mamo, as a potentially regulator of both CPs and melanin pathway genes, appear to be a potent modulator of epithelial traits. Claiming that there are inherent mutational biases behind this is unwarranted.

Response: Thank you very much for your careful work. I completely agree with your statement that the genome exhibits a certain degree of programmability. On the one hand, some transcription factors can precisely control the spatiotemporal expression levels of some structural genes (such as pigment synthesis genes). On the other hand, these transcription factors are also subject to strict expression regulation. Because the color pattern is complex, changes in single or minority structural genes result in incomplete or imprecise changes in coloring patterns. Nevertheless, several regulatory factors can regulate multiple downstream target genes. Changes in their expression patterns can lead to holistic and significant changes in color patterns. There are long intergenic regions upstream of many important transcription factors, dozens of kilobase pairs (Kb) to hundreds of Kb, which may contain many different regulatory elements for better control of their expression patterns. Therefore, gene regulatory networks can directly regulate transcription factors to modulate a given type of trait. Transcription factors and their downstream target genes can form a functional module, which is similar to a functional module in software or operating systems. This regulation of transcription factors is simpler in terms of steps, which are similar to a single click switch button. The gene regulatory network regulates these modules in response to environmental changes and is widely recognized.

Some people do not agree that genetic variations can also be regulated. They claim that this is completely random. The infinite monkey theorem (Félix-Édouard-Justin-Émile Borel, 1909) states that if an infinite number of monkeys were given typewriters and an infinite amount of time, they would eventually produce the complete works of Shakespeare. Although this theory advocates randomness on the surface, its conclusions are full of inevitability (tail event). In nature, some things we observe do not have obvious regularity because they involve relatively complex factors, and the underlying logic is obscure and difficult to understand. We often name them random. However, as we gradually understand the logic behind this complex event, we can also recognize the procedural nature of this randomness.

Previously, chromosomal recombination during meiosis was believed to be a random event. However, currently, it is believed that the process is procedural. The occurrence of meiotic recombination mentioned earlier indicates that the genome has the ability to self-set the position of double-strand breaks to form new allelic forms. Because meiotic recombination is programmed, transcription factors that recognize DNA sites, enzymes that cleave double strands, and DNA repair systems exist, programming can also introduce genetic variation. A study in plants has provided insights into this programmed mutation (Monroe JG, 2023, nature). Frequent changes in the expression patterns of some transcription factors occur between and/or within species. In this article, we only discuss the possible reasons for variations in the expression patterns of some transcription factors in a general manner and simple reasoning. We have added an analysis of the response of wild silkworms and improved the relevance of the discussion.

Monroe JG, Srikant T, Carbonell-Bejerano P, Becker C, Lensink M, Exposito-Alonso M, Klein M, Hildebrandt J, Neumann M, Kliebenstein D, Weng ML, Imbert E, Ågren J, Rutter MT, Fenster CB, Weigel D. Mutation bias reflects natural selection in Arabidopsis thaliana. Nature. 2022 Feb;602(7895):101-105. doi: 10.1038/s41586-021-04269-6. Epub 2022 Jan 12. Erratum in: Nature. 2023 Aug;620(7973):

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Please structure your Discussion with section headers.

Response: Thank you very much for your careful work. We have added relevant section headers.

As explained in my public review, I found the two last sections of the Discussion to be dispersed and confusing. I also must say that I carefully read the Response to Reviewers on this, which helped me to better understand the authors' intentions here. Please consider the revision of this Discussion as this feels extremely speculative difficult to connect with Bm-mamo.

Response: Thank you very much for your careful work. We have rewritten this part of the content.

typo: were found near the TTS of yellow --> TSS

Response: Thank you very much for your careful work. We have made these modifications.

l. 234 :"expression level of the 18 CP genes in the integument". Consider adding a mention of Figure 7 here, as only Fig. S10 is cited here.

Response: Thank you very much for your careful work. We have made these modifications.

Editorial comment on the second half of the Abstract:

Wu et al : "We found that Bm-mamo can comprehensively regulate the expression of related pigment synthesis and cuticular protein genes to form color patterns. This indicates that insects have a genetic basis for coordinate regulation of the structure and shape of the cuticle, as well as color patterns. This genetic basis provides the possibility for constructing the complex appearances of some insects. This study provides new insight into the regulation of color patterns."

I respectfully suggest a more accurate rephrasing, where the methods are mentioned, and where the logical argument is more straightforward. For example

"Using RNAi and CRISPR we show that Bm-mamo is a repressor or dark melanin patterns in the larval epithelium. Using in-vitro binding assays and gene expression profiling in wild-type and mutant larvae, we also show that Bm-mamo likely regulate the expression of related pigment synthesis and cuticular protein genes in a coordinated manner to mediate its role in color pattern formation. This mechanism is consistent with a dual role of this transcription factor in regulating both the structure and shape of the cuticle and pigments that are embedded within it. This study provides new insight into the regulation of color patterns as well as in the construction more complex epithelial features in some insects."

I hope this let the ideas of the original version transpire as the authors intended.

Response: Thank you very much for your careful work. We have made these modifications.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The color pattern of insects is one of the most dazzling adaptive evolutionary phenotypes. However, the molecular regulation of this color pattern is not fully understood. In this study, we found that the transcription factor Bm-mamo is responsible for black dilute (bd) allele mutations in the silkworm. Bm-mamo belongs to the BTB zinc finger family and is orthologous to mamo in Drosophila melanogaster. This gene has a conserved function in gamete production in Drosophila and silkworms and evolved a pleiotropic function in the regulation of color patterns in caterpillars. We found that Bm-mamo can comprehensively regulate the expression of related pigment synthesis and cuticular protein genes to form color patterns. This indicates that insects have a genetic basis for coordinate regulation of the structure and shape of the cuticle, as well as color patterns. This genetic basis provides the possibility for constructing the complex appearances of some insects. This study provides new insight into the regulation of color patterns.

Revised version: The color pattern of insects is one of the most diverse adaptive evolutionary phenotypes. However, the molecular regulation of this color pattern is not fully understood. In this study, we found that the transcription factor Bm-mamo is responsible for black dilute (bd) allele mutations in the silkworm. Bm-mamo belongs to the BTB zinc finger family and is orthologous to mamo in Drosophila melanogaster. This gene has a conserved function in gamete production in Drosophila and silkworms and has evolved a pleiotropic function in the regulation of color patterns in caterpillars. Using RNAi and clustered regularly interspaced short palindromic repeats (CRISPR) technology, we showed that Bm-mamo is a repressor or has dark melanin patterns in the larval epidermis. Using in vitro binding assays and gene expression profiling in wild-type and mutant larvae, we also showed that Bm-mamo likely regulates the expression of related pigment synthesis and cuticular protein genes in a coordinated manner to mediate its role in color pattern formation. This mechanism is consistent with the dual role of this transcription factor in regulating both the structure and shape of the cuticle and the pigments that are embedded within it. This study provides new insight into the regulation of color patterns as well as into the construction of more complex epidermis features in some insects.

Introduction

Earlier version: Insects often display stunning colors, and the patterns of these colors have been shown to be involved in behavior(1); immunity(2); thermoregulation(3); UV protection(4); and especially visual antagonism of predators via aposematism(5), mimicry(6), cryptic color patterns(7) or some combination of the above (8). In addition, the color patterns are divergent and convergent among populations(9). Due to these striking visual features and highly active adaptive evolutionary phenotypes, the genetic basis and evolutionary mechanism of color patterns have been a topic of interest for a long time.

Revised version: Insects often display stunning colors, and the patterns of these colors have been shown to be involved in behavior(1), immunity(2), thermoregulation(3), and UV protection(4); in particular, visual antagonism of predators via aposematism(5), mimicry(6), cryptic color patterns(7) or some combination of the above (8). In addition, the color patterns are divergent and convergent among populations(9). Due to these striking visual features and highly active adaptive evolutionary phenotypes, the genetic basis and evolutionary mechanism of color patterns have long been a topic of interest.

Earlier version: Insect coloration can be pigmentary, structural, or bioluminescent. Pigments are mainly synthesized by the insects themselves and form solid particles that are deposited in the cuticle of the body surface and the scales of the wings(10, 11). Interestingly, recent studies have found that bile pigments and carotenoid pigments synthesized through biological synthesis are incorporated into body fluids and filled in the wing membranes of two butterflies (Siproeta stelenes and Philaethria diatonica) via hemolymph circulation, providing color in the form of liquid pigments(12). The pigments form colors by selective absorption and/or scattering of light depending on their physical properties(13). However, structural color refers to colors, such as metallic colors and iridescence, generated by optical interference and grating diffraction of the microstructure/nanostructure of the body surface or appendages (such as scales)(14, 15). Pigment color and structural color are widely distributed in insects and can only be observed by the naked eye in illuminated environments. However, some insects, such as fireflies, exhibit colors (green to orange) in the dark due to bioluminescence(16). Bioluminescence occurs when luciferase catalyzes the oxidation of small molecules of luciferin(17). In conclusion, the color patterns of insects have evolved to be highly sophisticated and are closely related to their living environments. For example, cryptic color can deceive animals via high similarity to the surrounding environment. However, the molecular mechanism by which insects form precise color patterns to match their living environment is still unknown.

Revised version: Insect coloration can be pigmentary, structural, or bioluminescent. Pigments are synthesized by insects themselves and form solid particles that are deposited within the cuticle of the body surface and the scales of the wings(10, 11). Interestingly, recent studies have shown that bile pigments and carotenoid pigments synthesized through biological synthesis are incorporated into body fluids and fill in the wing membranes of two butterflies (Siproeta stelenes and Philaethria diatonica) via hemolymph circulation, providing color in the form of liquid pigments(12). These pigments form colors by selective absorption and/or scattering of light depending on their physical properties(13). However, structural color refers to colors, such as metallic colors and iridescence, generated by optical interference and grating diffraction of the microstructure/nanostructure of the body surface or appendages (such as scales)(14, 15). Pigment color and structural color are widely distributed in insects and can only be observed by the naked eye in illuminated environments. However, some insects, such as fireflies, exhibit colors (green to orange) in the dark due to bioluminescence(16). Bioluminescence occurs when luciferase catalyzes the oxidation of small molecules of luciferin(17). In conclusion, the color patterns of insects have evolved to become highly sophisticated and are closely related to their living environment. For example, cryptic color can deceive animals via high similarity to the surrounding environment. However, the molecular mechanism by which insects form precise color patterns to match their living environment has not been determined.

Earlier version: Recent research has identified the metabolic pathways of related pigments, such as melanins, pterins, and ommochromes, in Lepidoptera(18). A deficiency of enzymes in the pigment metabolism pathway can lead to changes in color patterns(19, 20). In addition to pigment synthesis, the microstructure/nanostructure of the body surface and wing scales are important factors that influence body color patterns. The body surface and wing scales of Lepidoptera are mainly composed of cuticular proteins (CPs) (21, 22). There are multiple genes encoding CPs in Lepidopteran genomes. For example, in Bombyx mori, over 220 genes encode CPs(23). However, there are still many unknowns regarding the function of CPs and the molecular mechanisms underlying their fine-scale localization in the cuticle.

Revised version: Recent research has identified the metabolic pathways associated with related pigments, such as melanins, pterins, and ommochromes, in Lepidoptera(18). A deficiency of enzymes in the pigment metabolism pathway can lead to changes in color patterns(19, 20). In addition to pigment synthesis, the microstructure/nanostructure of the body surface and wing scales are important factors that influence body color patterns. The body surface and wing scales of Lepidoptera are composed mainly of cuticular proteins (CPs)(21, 22). There are multiple genes encoding CPs in Lepidopteran genomes. For example, in Bombyx mori, more than 220 genes encode CPs(23). However, the functions of CPs and the molecular mechanisms underlying their fine-scale localization in the cuticle are still unclear.

Earlier version: In addition, some pleiotropic factors, such as wnt1(24), Apontic-like(25), clawless(26), abdominal-A(27), abdominal-B(28), engrailed(29), antennapedia(30), optix(31), bric à brac (bab)(32), and Distal-less (Dll)(33), play an important role in the regulation of color patterns. The molecular mechanism by which these factors participate in the regulation of body color patterns needs further study. In addition, there are many undiscovered factors in the gene regulatory network (GRN) involved in the regulation of insect color patterns. The identification of new color pattern regulatory genes and the study of their molecular mechanisms would be helpful for further understanding color patterns.

Revised version: In addition, several pleiotropic factors, such as wnt1(24), Apontic-like(25), clawless(26), abdominal-A(27), abdominal-B(28), engrailed(29), antennapedia(30), optix(31), bric à brac (bab)(32), and Distal-less (Dll)(33), play important roles in the regulation of color patterns. The molecular mechanism by which these factors participate in the regulation of body color patterns needs further study. In addition, there are many undiscovered factors in the gene regulatory network (GRN) involved in the regulation of insect color patterns. The identification of new color pattern regulatory genes and the study of their molecular mechanisms would be helpful for further understanding color patterns.

Earlier version: Silkworms (B. mori) have been completely domesticated for more than 5000 years and are famous for their silk fiber production(34). Because of this long-term domestication, artificial selection and genetic research, its genetic background is well resolved, and approximately 200 strains of mutants with variable color patterns have been preserved(35), which provide a good resource for color pattern research. The black dilute (bd) mutant, which exhibits recessive Mendelian inheritance, has a dark gray larval body color, and the female is sterile. bdf (black dilute fertile) is one of the alleles, leading to a lighter body color than that of bd and fertile females (Fig. 1). The two mutations were mapped to a single locus at 22.9 centimorgans (cM) of linkage group 9(36). No pigmentation-related genes have been reported at this locus. Thus, research may reveal a new color pattern-related gene, which stimulated our interest.

Revised version: Silkworms (B. mori) have been completely domesticated for more than 5000 years and are famous for their silk fiber production(34). Because of this long history of domestication, artificial selection and genetic research, its genetic background has been well studied, and approximately 200 strains of mutants with variable color patterns have been preserved(35), which provides a good resource for color pattern research. The black dilute (bd) mutant, which exhibits recessive Mendelian inheritance, has a dark gray larval body color, and the female is sterile. Black dilute fertile (bdf) is an allele that leads to a lighter body color than bd and fertile females (Fig. 1). The two mutations were mapped to a single locus at 22.9 centimorgans (cM) of linkage group 9(36). No pigmentation-related genes have been reported at this locus. Thus, this research may reveal a new color pattern-related gene, which stimulated our interest.

Phenotypes of bd, bdf, and wild-type Dazao larvae.

The epidermis of bd is dark gray, and the epidermis of bdf is light gray at the 5th instar (day 3) of silkworm larvae. The bar indicates 1 centimeter.

Results

Candidate gene of the bd allele

Earlier version: To identify the genomic region responsible for the bd alleles, positional cloning of the bd locus was performed. Due to the female infertility of the bd mutant and the fertility of females with the bdf allele, we used bdf and the wild-type Dazao strain as the parents for mapping analysis. The 1162 back-crossed filial 1st (BC1M) generation individuals from bdfand Dazao were used for fine mapping with molecular markers (Fig. 2). A genomic region of approximately 390 kilobases (kb) was responsible for the bd phenotype (Table S1). In the SilkDB database(37), this region included five predicted genes (BGIBMGA012516, BGIBMGA012517, BGIBMGA012518, BGIBMGA012519 and BGIBMGA014089). In addition, we analyzed the predictive genes for this genomic region from the GenBank(38) and SilkBase(39) databases (Fig. S1). The number of predicted genes varied among different databases. We performed sequence alignment analysis of the predicted genes in the three databases to determine their correspondence. Then, Real-time Quantitative polymerase chain reaction (qPCR) was performed, which showed that BGIBMGA012517 and BGIBMGA012518 were significantly downregulated on Day 3 of the fifth instar of larvae in the bd phenotype individuals, while there was no difference in the expression levels of other genes (Fig. S1). These two genes were predicted as a single locus (LOC101738295) in GenBank. To determine the gene structure of BGIBMGA012517 and BGIBMGA012518, we used forward primers for the BGIBMGA012517 gene and reverse primers for the BGIBMGA012518 gene to amplify cDNA from the wild-type Dazao strain. By gene cloning, the predicted genes BGIBMGA012517 and BGIBMGA012518 were proven to be one gene. For the convenience of description, we have temporarily called it the 12517/8 gene.

Revised version: To identify the genomic region responsible for the bd alleles, positional cloning of the bd locus was performed. Due to the female infertility of the bd mutant and the fertility of females with the bdf allele, we used bdf and the wild-type Dazao strain as parents for mapping analysis. The 1162 back-crossed filial 1st (BC1M) generation individuals from bdfand Dazao were subjected to fine mapping with molecular markers (Fig. 2). A genomic region of approximately 390 kilobases (kb) was responsible for the bd phenotype (Table S1). According to the SilkDB database(37), this region included five predicted genes (BGIBMGA012516, BGIBMGA012517, BGIBMGA012518, BGIBMGA012519 and BGIBMGA014089). In addition, we analyzed the predictive genes for this genomic region from the GenBank(38) and SilkBase(39) databases (Fig. S1). The number of predicted genes varied among the different databases. We performed sequence alignment analysis of the predicted genes in the three databases to determine their correspondence. Real-time quantitative polymerase chain reaction (qPCR) was subsequently performed, which revealed that BGIBMGA012517 and BGIBMGA012518 were significantly downregulated on Day 3 of the fifth instar of the larvae in the bd phenotype individuals, while there was no difference in the expression levels of the other genes (Fig. S1). These two genes were predicted to be associated with a single locus (LOC101738295) in GenBank. To determine the gene structures of BGIBMGA012517 and BGIBMGA012518, we used forward primers for the BGIBMGA012517 gene and reverse primers for the BGIBMGA012518 gene to amplify cDNA from the wild-type Dazao strain. By gene cloning, the predicted genes BGIBMGA012517 and BGIBMGA012518 were proven to be one such gene. For convenience, we temporarily called this the 12517/8 gene.

Positional cloning of the bd locus.

(A) We used 532 BC1 individuals to map the bd locus between PCR markers c and i. The numbers above the DNA markers indicate recombination events.

(B) A total of 1162 BC1 individuals were used to narrow the bd locus to an approximately 400 kb genomic region.

(C) Partial enlarged view of the region responsible for bd. This region contains 4 predicted genes, BGIBMGA012517, BGIBMGA012518, BGIBMGA012519 and BGIBMGA014089.

(D) Analysis of nucleotide differences in the region responsible for bd. The green block indicates the deletion of the genome in bd mutants. The black vertical lines indicate the SNPs and indels of bdf mutants.

Earlier version: The 12517/8 gene produces two transcripts; the open reading frame (ORF) of the long transcript is 2397 bp, and the ORF of the short transcript is 1824 bp, which shares the same 5’-terminus as the wild-type Dazao strain (Fig. S2). The 12517/8 gene showed significantly lower expression in the bdf mutant (Fig. S3), and multiple variations were found in the nearby region of this gene by comparative genomic analysis between bdf and Dazao (Table S2). In addition, 12517/8 was completely silenced due to the deletion of DNA fragments (approximately 168 kb) from the first upstream intron and an insertion of 3560 bp in the bd mutant (Fig. S4).

Revised version: The 12517/8 gene produces two transcripts; the open reading frame (ORF) of the long transcript is 2397 bp, and the ORF of the short transcript is 1824 bp, which is the same 5’-terminus as that of the wild-type Dazao strain (Fig. S2). The 12517/8 gene showed significantly lower expression in the bdf mutant (Fig. S3), and multiple variations were found in the region nearby this gene by comparative genomic analysis between bdf and Dazao (Table S2). In addition, 12517/8 was completely silenced due to the deletion of DNA fragments (approximately 168 kb) from the first upstream intron and an insertion of 3560 bp in the bd mutant (Fig. S4).

Earlier version: To predict the function of the 12517/8 gene, we performed a BLAST search using its full-length sequence and found a transcription factor, the maternal gene required for meiosis (mamo) in D. melanogaster, that had a high sequence similarity to that of the 12517/8 gene (Fig. S5). Therefore, we named the 12517/18 gene Bm-mamo; the long transcript was designated as Bm-mamo-L, and the short transcript was designated as Bm-mamo-S.

Revised version: To predict the function of the 12517/8 gene, we performed a BLAST search using its full-length sequence and found a transcription factor, the maternal gene required for meiosis (mamo) in D. melanogaster, that had high sequence similarity to that of the 12517/8 gene (Fig. S5). Therefore, we named the 12517/18 gene Bm-mamo; the long transcript was designated Bm-mamo-L, and the short transcript was designated Bm-mamo-S.

Earlier version: The mamo gene belongs to the Broad-complex, Tramtrack and bric à brac/poxvirus zinc finger protein (BTB-ZF) family. In the BTB-ZF family, the zinc finger serves as a recognition motif for DNA-specific sequences, and the BTB domain promotes oligomerization and the recruitment of other regulatory factors(40). Most of these factors are typically transcription repressors, such as nuclear receptor corepressor (N-CoR) and silencing mediator for retinoid and thyroid hormone receptor (SMRT)(41), but some are activators, such as p300(42). Therefore, these features commonly serve as regulators of gene expression. mamo is enriched in embryonic primordial germ cells (PGCs) in D. melanogaster. Individuals deficient in mamo are able to undergo oogenesis but fail to execute meiosis properly, leading to female infertility in D. melanogaster(43). Bm-mamo was identified as an important candidate gene for further analysis.

Revised version: The mamo gene belongs to the Broad-complex, Tramtrack and bric à brac/poxvirus zinc finger protein (BTB-ZF) family. In the BTB-ZF family, the zinc finger serves as a recognition motif for DNA-specific sequences, and the BTB domain promotes oligomerization and the recruitment of other regulatory factors(40). Most of these factors are transcription repressors, such as nuclear receptor corepressor (N-CoR) and silencing mediator for retinoid and thyroid hormone receptor (SMRT)(41), but some are activators, such as p300(42). Therefore, these features commonly serve as regulators of gene expression. mamo is enriched in embryonic primordial germ cells (PGCs) in D. melanogaster. Individuals deficient in mamo are able to undergo oogenesis but fail to execute meiosis properly, leading to female infertility in D. melanogaster(43). Bm-mamo was identified as an important candidate gene for further analysis.

Expression pattern analysis of Bm-mamo

Earlier version: To analyze the expression profiles of Bm-mamo, we performed quantitative PCR. The expression levels of the Bm-mamo gene were investigated in the whole body at different developmental stages, from the embryonic stage to the adult stage, in the Dazao strain. The gene was highly expressed in the molting stage of caterpillars, and its expression was upregulated in the later pupal and adult stages (Fig. 3A). This suggests that the Bm-mamo gene responds to the ecdysone titer and participates in the processes of molting and metamorphosis in silkworms. In the investigation of tissue-specific expression levels in the 5th-instar 3rd-day larvae of the Dazao strain, the midgut, head, and epidermis had high expression levels; the trachea, nerves, silk glands, testis, ovary, muscle, wing disc and Malpighian tubules had moderate expression levels; and the blood and fat bodies had low expression levels (Fig. 3B). This suggests that Bm-mamo is involved in the developmental regulation of multiple silkworm tissues. Due to the melanism of the epidermis of the bd mutant and the high expression level of the Bm-mamo gene in the epidermis, we measured the expression level of this gene in the epidermis of the 4th to 5th instar of the Dazao strain. In the epidermis, the Bm-mamo gene was upregulated in the molting period, and the highest expression was observed at the beginning of molting (Fig. 3C).

Revised version: To analyze the expression profiles of Bm-mamo, we performed quantitative PCR. The expression levels of the Bm-mamo gene in the Dazao strain were investigated throughout the body at different developmental stages, from the embryonic stage to the adult stage. The gene was highly expressed in the molting stage of caterpillars, and its expression was upregulated in the later pupal and adult stages (Fig. 3A). This finding suggested that the Bm-mamo gene responds to ecdysone and participates in the processes of molting and metamorphosis in silkworms.

Spatiotemporal expression of Bm-mamo.

Earlier version: (A) Temporal expression of Bm-mamo. In the molting stage and adult stage, this gene is significantly upregulated. M: molting stage, W: wandering stage, P: pupal stage, P1: Day 1 of pupal stage, A: adult stage, 4th3d indicates 4th instar day 3. 1st to 5th denote the first instar of larvae to fifth instar of larvae, respectively.

Revised version: (A) Temporal expression of Bm-mamo. In the molting stage and adult stage, this gene was significantly upregulated. M: molting stage, W: wandering stage, P: pupal stage, P1: Day 1 of the pupal stage, A: adult stage, 4th3d indicates the 4th instar on Day 3. 1st to 5th denote the first instar of larvae to fifth instar of larvae, respectively.

(B) Tissue-specific expression of 4th-instar molting larvae. Bm-mamo had relatively high expression levels in the midgut, head, and epidermis. AMSG: anterior division of silk gland and middle silk gland, PSG: posterior silk gland.

(C) Detailed analysis of Bm-mamo at the 4th larval stage in the epidermis of the Dazao strain. Bm-mamo expression is upregulated during the molting stage. HCS indicates the head capsule stage. The “h” indicates the hour, 90 h: at 90 hours of the 4th instar.

Added in the revised version: According to the investigation of tissue-specific Dazao strain expression in 5th-instar 3rd-day larvae, the midgut, head, and epidermis exhibited high expression levels; the trachea, nerves, silk glands, testis, ovary, muscle, wing disc and Malpighian tubules exhibited moderate expression levels; and the blood and fat bodies exhibited low expression levels (Fig. 3B). These findings suggested that Bm-mamo is involved in the developmental regulation of multiple silkworm tissues. Due to the melanism of the epidermis of the bd mutant and the high expression level of the Bm-mamo gene in the epidermis, we measured the expression level of this gene in the epidermis of the 4th to 5th instars of the Dazao strain. In the epidermis, the Bm-mamo gene was upregulated during the molting period, and the highest expression was observed at the beginning of molting (Fig. 3C).

Functional analyses of Bm-mamo

Earlier version: To study the function of the Bm-mamo gene, we carried out an RNA interference (RNAi) experiment. Short interfering RNA (siRNA) was injected into the hemolymph of silkworms, and an electroporation experiment(44) was immediately conducted. We found significant melanin pigmentation in the epidermis of the newly molted 5th-instar larvae. This indicates that Bm-mamo deficiency can cause melanin pigmentation (Fig. 4). The melanism phenotype of RNAi individuals was similar to that of bd mutants.

Revised version: To study the function of the Bm-mamo gene, we carried out an RNA interference (RNAi) experiment. Short interfering RNA (siRNA) was injected into the hemolymph of silkworms, and an electroporation experiment(44) was immediately conducted. We found significant melanin pigmentation in the epidermis of the newly molted 5th-instar larvae. These findings indicate that Bm-mamo deficiency can cause melanin pigmentation (Fig. 4). The melanistic phenotype of the RNAi individuals was similar to that of the bd mutants.

Earlier version: (A) siRNA was introduced by microinjection followed by electroporation. “+” and “-” indicate the positive and negative poles of the electrical current. Scale bars, 1 cm.

Revised version: (A) siRNA was introduced by microinjection followed by electroporation. “+” and “-” indicate the positive and negative poles of the electrical current, respectively. Scale bars, 1 cm.

(B) Partial magnification of the siRNA experimental group and negative control group. Scale bars, 0.2 cm.

(C) and (D) Relative expression levels of Bm-mamo in the negative control and RNAi groups were determined by qPCR analysis. The means ± s.d.s. Numbers of samples are shown in the upper right in each graph. **P<0.01, paired Student’s t test (NS, not significant).

Earlier version: (E) Efficiency statistics of RNAi.

Revised version: (E) Statistical analysis of the efficiency of the RNAi.

Earlier version: In addition, gene knockout was performed. Ribonucleoproteins (RNPs) generated by guide RNA (gRNA) and recombinant Cas9 protein were injected into 450 silkworm embryos. In the G0 generation, individuals with a mosaic melanization phenotype were found. These melanistic individuals were raised to moths and then crossed. The homozygous line with gene knockout was obtained through generations of screening. The gene-edited individuals had a significantly melanistic body color, and the female moths were sterile (Fig. 5).

Revised version: In addition, gene knockout was performed. Ribonucleoproteins (RNPs) generated from guide RNA (gRNA) and recombinant Cas9 protein were injected into 450 silkworm embryos. In the G0 generation, individuals with a mosaic melanization phenotype were found. These melanistic individuals were raised to moths and subsequently crossed. A homozygous line with a gene knockout was obtained through generations of screening. The gene-edited individuals had a significantly melanistic body color, and the female moths were sterile (Fig. 5).

Earlier version: Knockout of Bm-mamo

Revised version: Bm-mamo knockout

Earlier version: (A) Larval phenotype in G3 fourth-instar larvae of Dazao targeted for Bm-mamo. Scale bars, 1 cm.

Revised version: (A) Larval phenotype of G3 fourth-instar larvae of Dazao targeted for Bm-mamo. Scale bars, 1 cm.

(B) Partial magnification of the Bm-mamo knockout individual and control. Scale bars, 0.2 cm.

Earlier version: (C) After 48 hours of laying eggs, the pigmentation of the eggs indicates that those produced by knockout homozygous females cannot undergo normal pigmentation and development.

Revised version: (C) After 48 hours of egg laying, the pigmentation of the eggs indicates that those produced by knockout homozygous females cannot undergo normal pigmentation and development.

Earlier version: (D) Genomic structure of Bm-mamo. Open reading frame (blue), untranslated region (black). The gRNA 1 and gRNA 2 sequences are shown.

Revised version: (D) Genomic structure of Bm-mamo. The open reading frame (blue) and untranslated region (black) are shown. The gRNA 1 and gRNA 2 sequences are shown.

(E) Sequences of the Bm-mamo knockout individuals. Lines 1 and 2 indicate deletions of 15 and 71 bp, respectively.

These results indicated that the Bm-mamo gene negatively regulates melanin pigmentation in caterpillars and participates in the reproductive regulation of female moths.

Downstream target gene analysis

Earlier version: Bm-mamo belongs to the zinc finger protein family, which specifically recognizes downstream DNA sequences according to their zinc fingers. We identified homologous genes of mamo in multiple species and conducted a phylogenetic analysis (Fig. S6). Sequence alignment revealed that the amino acid residues of the zinc finger motif of the mamo-S protein were highly conserved among 57 species. The insufficient prediction of alternative transcripts of mamo in different species in GenBank may have resulted in the mamo-L protein being found in only 30 species (Fig. S7). As the mamo protein carries a tandem Cys2His2 zinc finger (C2H2-ZF) motif, it can directly bind to DNA sequences. Previous research has suggested that the ZF-DNA binding interface can be understood as a “canonical binding model”, in which each finger contacts DNA in an antiparallel manner. The binding sequence of the C2H2-ZF motif is determined by the amino acid residue sequence of its α-helical component. Considering the first amino acid residue in the α-helical region of the C2H2-ZF domain as position 1, positions –1, 2, 3, and 6 are key amino acids in the recognition and binding of DNA. The residues at positions –1, 3, and 6 specifically interact with base 3, base 2, and base 1 of the DNA sense sequence, respectively, while the residue at position 2 interacts with the complementary DNA strand(45, 46). To analyze the downstream target genes of mamo, we first predicted its DNA binding motifs using online software (http://zf.princeton.edu) based on the canonical binding model(47). In addition, the DNA-binding sequence of mamo (TGCGT) in Drosophila, confirmed by electrophoretic mobility shift assay (EMSA) experiments(48), has a consensus sequence with the predicted binding site sequence of Bm-mamo-S (GTGCGTGGC), and the predicted sequence was longer. This indicates that the predicted results of the DNA-binding site have good reliability.

Revised version: Bm-mamo belongs to the zinc finger protein family, which specifically recognizes downstream DNA sequences according to their zinc fingers. We identified homologous genes of mamo in multiple species and conducted a phylogenetic analysis (Fig. S6). Sequence alignment revealed that the amino acid residues of the zinc finger motif of the mamo-S protein were highly conserved among 57 species. The insufficient ability of the different species in GenBank to predict alternative transcripts of mamo may have resulted in the mamo-L protein being found in only 30 species (Fig. S7). As the mamo protein contains a tandem Cys2His2 zinc finger (C2H2-ZF) motif, it can directly bind to DNA sequences. Previous research has suggested that the ZF-DNA binding interface can be understood as a “canonical binding model”, in which each finger contacts DNA in an antiparallel manner. The binding sequence of the C2H2-ZF motif is determined by the amino acid residue sequence of its α-helical component. The first amino acid residue in the α-helical region of the C2H2-ZF domain is at position 1, and positions −1, 2, 3, and 6 are key amino acids involved in the recognition and binding of DNA. The residues at positions −1, 3, and 6 specifically interact with base 3, base 2, and base 1 of the DNA sense sequence, respectively, while the residue at position 2 interacts with the complementary DNA strand(45, 46). To analyze the downstream target genes of mamo, we first predicted the DNA-binding motifs of these genes using online software (http://zf.princeton.edu) based on the canonical binding model(47). In addition, the DNA-binding sequence of mamo (TGCGT) in Drosophila was confirmed by electrophoretic mobility shift assay (EMSA)(48), which has a consensus sequence with the predicted binding site sequence of Bm-mamo-S (GTGCGTGGC), and the predicted sequence was longer. This indicates that the predicted results for the DNA-binding site have good reliability. Furthermore, the predicted DNA-binding sites of Bm-mamo-L and Bm-mamo-S were highly consistent with those of mamo orthologs in different species (Fig. S8). This finding suggested that the protein may regulate similar target genes between species.

Removed in the revised version: Furthermore, the predicted DNA-binding sites of Bm-mamo-L and Bm-mamo-S were highly consistent with those of mamo orthologs in different species (Fig. S8). This suggests that the protein may regulate similar target genes between species.

Earlier version: C2H2-ZF transcription factors function by recognizing and binding to cis-regulatory sequences in the genome, which harbor cis-regulatory elements (CREs)(49). CREs are broadly classified as promoters, enhancers, silencers and insulators(50). CREs are often located near their target genes, such as enhancers, which are typically located upstream (5’) or downstream (3’) of the gene they regulate, or in introns, but approximately 12% of CREs are located far from their target gene(51). Therefore, we first investigated the genome range 2 kb upstream and downstream of the predicted genes in silkworms.

Revised version: C2H2-ZF transcription factors function by recognizing and binding to cis-regulatory sequences in the genome, which harbor cis-regulatory elements (CREs)(49). CREs are broadly classified as promoters, enhancers, silencers or insulators(50). CREs are often near their target genes, such as enhancers, which are typically located upstream (5’) or downstream (3’) of the gene they regulate or in introns, but approximately 12% of CREs are located far from their target gene(51). Therefore, we first investigated the 2 kb upstream and downstream regions of the predicted genes in silkworms.

Earlier version: The predicted position weight matrices (PWMs) of the recognized sequences of Bm-mamo and the Find Individual Motif Occurrences (FIMO) software of MEME were used to perform silkworm whole-genome scanning for possible downstream target genes. The genomic regions 2 kb upstream and downstream of the 14,623 predicted genes in silkworms were investigated. A total of 10,622 genes contained the recognition sites of the Bm-mamo protein in the silkworm genome (Fig. S9).

Revised version: The predicted position weight matrices (PWMs) of the recognized sequences of Bm-mamo protein and the Find Individual Motif Occurrences (FIMO) software of MEME were used to perform silkworm whole-genome scanning for possible downstream target genes. The 2 kb upstream and downstream regions of the 14,623 predicted genes in silkworms were investigated. A total of 10,622 genes contained recognition sites within 2 kb of the upstream/downstream region of the Bm-mamo protein in the silkworm genome (Fig. S9, Table S3 & Table S4).

Earlier version: Moreover, we compared transcriptome data of the integument tissue between homozygotes and heterozygotes of the bd mutant at the 4th instar/beginning molting stage(52). In the integument tissue, 10,072 genes (∼69% of the total predicted genes of silkworm) were expressed in heterozygotes, and 9,853 genes (∼67% of the total predicted genes) were expressed in homozygotes of the bd mutant. In addition, there were 191 genes with significant expression differences between homozygotes (bd/bd) and heterozygotes (+/bd) by comparative transcriptome analysis (Table S3)(52). Protein functional annotation was performed, and 19 CP genes were significantly differentially expressed between heterozygotes and homozygotes of bd. In addition, the orthologs of these CPs were analyzed in Danaus plexippus, Papilio xuthu and D. melanogaster (Table S4). Furthermore, we identified 53 enzyme-encoding genes, 17 antimicrobial peptide genes, 6 transporter genes, 5 transcription factor genes, 5 cytochrome genes, and others. CP genes were significantly enriched among the differentially expressed genes (DEGs), and previous studies have found that CPs can affect pigmentation(53). Therefore, we first conducted an investigation of the CP genes. Among them, 18 CP genes had Bm-mamo binding sites in the upstream and downstream 2 kb genomic regions. In addition, we investigated the expression level of the 18 CP genes in the integument from the 4th instar (day 1) to the beginning of the 5th instar of Bm-mamo knockout lines (Fig. S10). The CP genes were significantly upregulated at one or several time points in homozygous (mamo-/mamo-) individuals compared with heterozygous (mamo-/+) Bm-mamo knockout line individuals.

Revised version: Moreover, we compared the transcriptome data of integument tissue between homozygotes and heterozygotes of the bd mutant at the 4th instar/beginning molting stage(52). In the integument tissue, 10,072 genes (∼69% of the total predicted genes of silkworm) were expressed in heterozygotes, and 9,853 genes (∼67% of the total predicted genes) were expressed in homozygotes of the bd mutant. In addition, there were 191 genes whose expression significantly differed between homozygotes (bd/bd) and heterozygotes (+/bd) according to comparative transcriptome analysis (Table S5)(52). Protein functional annotation was performed, and 19 CP genes were significantly differentially expressed between heterozygotes and homozygotes of bd. In addition, the orthologs of these CPs were analyzed in Danaus plexippus, Papilio xuthu and D. melanogaster (Table S6). Furthermore, we identified 53 enzyme-encoding genes, 17 antimicrobial peptide genes, 6 transporter genes, 5 transcription factor genes, 5 cytochrome genes, and others. Among the differentially expressed genes (DEGs), CP was significantly enriched, and previous studies have shown that CPs can affect pigmentation(53). Therefore, we first investigated the expression of the CP genes. Among them, 18 CP genes had Bm-mamo binding sites in the upstream and downstream 2 kb genomic regions. In addition, we investigated the expression levels of the 18 CP genes in the integument from the 4th instar (Day 1) to the beginning of the 5th instar in the Bm-mamo knockout lines (Fig. 7& Fig. S10). The CP genes were significantly upregulated at one or several time points in homozygous (mamo-/mamo-) individuals compared with heterozygous (mamo-/+) Bm-mamo knockout individuals. Interestingly, the expression of the CP gene BmorCPH24 was significantly upregulated at the feeding stage in Bm-mamo knockout homozygotes. Previous studies have shown that BmorCPH24 deficiency can lead to a marked decrease in pigmentation in silkworm larvae(53). Therefore, the expression of some CP genes may be necessary for determining the color patterns of caterpillars.

Removed in the revised version: Interestingly, the CP gene BmorCPH24 was significantly upregulated at the feeding stage in Bm-mamo knockout homozygotes. Previous studies have shown that BmorCPH24 deficiency can lead to a marked decrease in pigmentation in silkworm larvae(53). Therefore, the expression of some CP genes may be necessary for color patterns in caterpillars.

Earlier version: In addition, the synthesis of pigment substances is an important driver of color patterns(54). Eight key genes (TH, DDC, aaNAT, ebony, black, tan, yellow and laccase2) involved in melanin synthesis(11) were investigated in heterozygous and homozygous gene knockout lines of Bm-mamo. Among them, yellow, tan, and DDC were significantly upregulated during the molting period in the homozygous Bm-mamo knockout individuals (Fig. 6). The upstream and/or downstream 2 kb genomic regions of yellow, tan and DDC contain the binding site of the Bm-mamo protein. In addition, the expression of yellow, DDC and tan can promote the generation of melanin(55).

Revised version: In addition, the synthesis of pigment substances is an important driver of color patterns(54). Eight key genes (TH, DDC, aaNAT, ebony, black, tan, yellow and laccase2) involved in melanin synthesis(11) were investigated in heterozygous and homozygous Bm-mamo gene knockout lines. Among them, yellow, tan, and DDC were significantly upregulated during the molting period in the homozygous Bm-mamo knockout individuals (Fig. 6). The upstream and/or downstream 2 kb genomic regions of yellow, tan and DDC contain the binding site of the Bm-mamo protein. In addition, the expression of yellow, DDC and tan can promote the generation of melanin(55). To explore the interaction between the Bm-mamo protein and its binding sequence, EMSA was conducted. The binding site of Bm-mamo-S (CTGCGTGGT) was located approximately 70 bp upstream of the transcription initiation site of the Bm-yellow gene. The EMSA results showed that the Bm-mamo-S protein expressed in prokaryotes can bind to the CTGCGTGGT sequence in vitro (Fig. S11). This finding suggested that the Bm-mamo-S protein can bind to the upstream region of the Bm-yellow gene and regulate its transcription. Therefore, the Bm-mamo gene may control the color pattern of caterpillars by regulating key melanin synthesis genes and 18 CP genes.

Melanin metabolism pathway and qPCR of related genes.

(A) The melanin metabolism pathway. Blue indicates amino acids and catecholamines, red indicates the names of genes, and black indicates the names of enzymes.

Earlier version: (B) Relative expression levels of eight genes in the heterozygous Bm-mamo knockout group (blue) and homozygous Bm-mamo knockout group (red) as determined by qPCR analysis. The means ± s.d.s. *P<0.05, paired Student’s t test. 4th3d indicates 4th instar day 3, M indicates molting, and h indicates hours.

Revised version: (B) Relative expression levels of eight genes in the heterozygous Bm-mamo knockout group (blue) and homozygous Bm-mamo knockout group (red) determined via qPCR analysis. The means ± s.d.s. *P<0.05, paired Student’s t test. 4th3d indicates the 4th instar on Day 3, M indicates molting, and h indicates hours.

Removed in the revised version: To explore the interaction between the Bm-mamo protein and its binding sequence, EMSA experiments were conducted. A binding site of Bm-mamo-S (CTGCGTGGT) was located approximately 70 bp upstream of the transcription initiation site of the Bm-yellow gene. The EMSA experiment showed that the Bm-mamo-S protein expressed in prokaryotes can bind to the CTGCGTGGT sequence in vitro (Fig. S11). This suggests that the Bm-mamo-S protein can bind to the upstream region of the Bm-yellow gene and regulate its transcription.

Removed in the revised version: Therefore, the Bm-mamo gene may control the color pattern of caterpillars by regulating key melanin synthesis genes and 18 CP genes.

Discussion

Earlier version: Insects have evolved many important phenotypes during the process of adapting to the environment. Among these traits, color pattern is one of the most interesting. The biochemical metabolic pathways of pigments, such as melanin, ommochromes, and pteridines, have been identified in insects(18). However, the regulation of pigment metabolism-related genes and the processes of the transport and deposition of pigment substances are not clear. In this study, we discovered that the Bm-mamo gene negatively regulates melanin pigmentation in caterpillars. When this gene is deficient, the body color of silkworms exhibits substantial melanism, and changes in the expression of some melanin synthesis genes and some CP genes are also significant.

Revised version: Insects have evolved many important phenotypes during the process of adapting to the environment. Among these traits, the color pattern is one of the most interesting. The biochemical metabolic pathways of pigments, such as melanin, ommochromes, and pteridines, have been identified in insects(18). However, the regulation of pigment metabolism-related genes and the processes involved in the transport and deposition of pigment substances are unclear. In this study, we discovered that the Bm-mamo gene negatively regulates melanin pigmentation in caterpillars. When this gene is deficient, the body color of silkworms exhibits substantial melanism, and changes in the expression of some melanin synthesis genes and some CP genes are also significant.

Added in the revised version: Structural genes of melanin synthesis

Earlier version: Deficiency of some genes in the melanin synthesis pathway, such as TH, yellow, ebony, tan, and aaNAT, can lead to variations in color patterns (11). However, they often lead only to localized pigmentation changes in later-instar caterpillars. For example, in yellow-mutant silkworms, the eye spot, lunar spot, star spot, spiracle plate, head, and some sclerotized areas appear reddish brown, and the other cuticle is consistent with that in the wild type in later instars of larvae (third instar to fifth instar)(56). This situation is highly similar to that of the tan mutant in silkworms(57). Why do the yellow and tan phenotypes appear only in a limited cuticle region in later instars of silkworm larvae? One possible reason is that the expression of yellow and tan is limited to a certain region by transcription regulators. Alternatively, other factors, such as CPs in the cuticle, may limit pigmentation by interacting with pigment substances in the cuticle.

Revised version: Deficiencies of some genes in the melanin synthesis pathway, such as TH, yellow, ebony, tan, and aaNAT, can lead to variations in color patterns (11). However, they often lead to localized pigmentation changes only in later-instar caterpillars. For example, in yellow-mutant silkworms, the eye spot, lunar spot, star spot, spiracle plate, head, and some sclerotized areas appear reddish brown, and the other cuticle is consistent with that in the wild type in later instars of larvae (third instar to fifth instar)(56). This situation is highly similar to that of the tan mutant in silkworms(57). Why do the yellow and tan phenotypes appear only in a limited cuticle region in later instars of silkworm larvae? One possible reason is that the expression of yellow and tan is limited to a certain region by transcription regulators. Alternatively, other factors, such as CPs in the cuticle, may limit pigmentation by interacting with pigment substances in the cuticle.

On the one hand, we investigated the expression levels of yellow and tan in the pigmented region (lunar spot) and nonpigmented region of the epidermis during the 4th molting of the wild-type Dazao strain. The expression level of yellow was significantly upregulated in the lunar spot (the epidermis on the dorsum of the fifth body segment) compared with the nonpigmented region (the epidermis on the dorsum of the sixth body segment) at 6 hours and 12 hours after molting.

Earlier version: Meanwhile, tan was significantly upregulated 18 hours after molting in the lunar spot region (Fig. S12). This suggests that the upregulated expression of pigment synthesis genes at key time points may be important for pigmentation. However, yellow and tan were still moderately expressed in the nonpigmented epidermis, although they did not cause significant melanin pigmentation. This indicates that pigment synthesis alone cannot dominantly determine the color pattern of the cuticle in caterpillars.

Revised version: Moreover, tan was significantly upregulated 18 hours after molting in the lunar spot region (Fig. S12). This finding suggested that the upregulated expression of pigment synthesis genes at key time points may be important for pigmentation. However, yellow and tan were still moderately expressed in the nonpigmented epidermis, although they did not cause significant melanin pigmentation. This finding indicates that pigment synthesis alone cannot determine the predominant color pattern of the cuticle in caterpillars.

Added in the revised version: Cuticular proteins participate in pigmentation

Removed in the revised version: On the other hand, synthesized pigment substances need to be transported from epidermal cells and embedded into the cuticle to allow pigmentation. Therefore, the correct cuticle structure and location of cuticular proteins in the cuticle may be important factors affecting pigmentation.

Earlier version: Previous studies have shown that the lack of expression of BmorCPH24, which encodes important components of the endocuticle, can lead to dramatic changes in body shape and a significant reduction in the pigmentation of caterpillars(53). We crossed Bo (BmorCPH24 null mutation) and bd to obtain F1(Bo/+Bo, bd/+), then self-crossed F1 and observed the phenotype of F2. The lunar spots and star spots decreased, and light-colored stripes appeared on the body segments, but the other areas still had significant melanin pigmentation in double mutation (Bo, bd) individuals (Fig. S13). However, in previous studies, introduction of Bo into L (ectopic expression of wnt1 results in lunar stripes generated on each body segment)(24) and U (overexpression of SoxD results in excessive melanin pigmentation of the epidermis)(58) strains by genetic crosses can remarkably reduce the pigmentation of L and U (53). Interestingly, there was a more significant decrease in pigmentation in the double mutants (Bo, L) and (Bo, U) than in (Bo, bd). This suggests that Bm-mamo has a stronger ability than wnt1 and SoxD to regulate pigmentation. On the one hand, mamo may be a stronger regulator of the melanin metabolic pathway, and on the other hand, mamo may regulate other CP genes to reduce the impact of BmorCPH24 deficiency.

Revised version: On the other hand, synthesized pigment substances need to be transported from epidermal cells and embedded into the cuticle to allow pigmentation. Therefore, the correct cuticle structure and location of cuticular proteins in the cuticle may be important factors affecting pigmentation. Previous studies have shown that a lack of expression of BmorCPH24, which encodes an important component of the endocuticle, can lead to dramatic changes in body shape and a significant reduction in the pigmentation of caterpillars(53). We crossed Bo (BmorCPH24 null mutation) and bd to obtain F1 (Bo/+Bo, bd/+) and then self-crossed F1 to observe the phenotype of F2. The area of lunar spots and star spots decreased, and light-colored stripes appeared on the body segments; however, the other areas still exhibited significant melanin pigmentation in double-mutation (Bo, bd) individuals (Fig. S13). However, in previous studies, the introduction of Bo into L (ectopic expression of wnt1 results in lunar stripes generated on each body segment)(24) and U (overexpression of SoxD results in excessive melanin pigmentation of the epidermis)(58) strains by genetic crosses markedly reduced the pigmentation of L and U(53). Interestingly, there was a more significant decrease in pigmentation in the double mutants (Bo, L) and (Bo, U) than in (Bo, bd). These findings suggested that Bm-mamo has a greater ability than does wnt1 and SoxD to regulate pigmentation. On the one hand, mamo may be a stronger regulator of the melanin metabolic pathway, and on the other hand, mamo may regulate other CP genes to reduce the impact of BmorCPH24 deficiency.

Earlier version: How do CPs affect pigmentation? One study showed that some CPs can form “pore canals” to transport some macromolecules(59). In addition, some CPs can be crosslinked with catecholamines, which are synthesized in the melanin metabolism pathway(60). Because there are no live cells in the cuticle, melanin precursor substances may be transported by the pore canals formed by some CPs and fixed to specific positions through cross-linking with CPs. The cuticular protein TcCPR4 is needed for the formation of pore canals of the cuticle in Tribolium castaneum(61). In contrast, the vertical pore canal is lacking in the less pigmented cuticles of T. castaneum(62). This suggests that the pore canals constructed by TcCPR4 may transport pigments and contribute to cuticle pigmentation in T. castaneum. Moreover, the melanin metabolites N-acetyldopamine (NADA) and N-β-alanyldopamine (NBAD) can target and sclerotize the cuticle by cross-linking with specific cuticular proteins (57). This suggests that pigments interact with specific CPs, thereby affecting pigmentation, hardening properties and the structure of the cuticle. Interestingly, a study found that pigments, in addition to absorbing specific wavelengths, can affect cuticle polymerization, density, and refractive index, which in turn affects the reflected wavelengths that produce structural color in butterfly wing scales(63).

Revised version: How do CPs affect pigmentation? One study showed that some CPs can form “pore canals” to transport macromolecules(59). In addition, some CPs can be crosslinked with catecholamines, which are synthesized in the melanin metabolism pathway(60). Because there are no live cells in the cuticle, melanin precursor substances may be transported by the pore canals formed by some CPs and fixed to specific positions through cross-linking with CPs. The cuticular protein TcCPR4 is needed for the formation of pore canals in the cuticle of Tribolium castaneum(61). In contrast, the vertical pore canal is lacking in the less pigmented cuticles of T. castaneum(62). This finding suggested that the pore canals constructed by TcCPR4 may transport pigments and contribute to cuticle pigmentation in T. castaneum. Moreover, the melanin metabolites N-acetyldopamine (NADA) and N-β-alanyldopamine (NBAD) can target and sclerotize the cuticle by cross-linking with specific cuticular proteins (57). This finding suggested that pigments interact with specific CPs, thereby affecting pigmentation, hardening properties and the structure of the cuticle. Interestingly, a study showed that in addition to absorbing specific wavelengths, pigments can affect cuticle polymerization, density, and the refractive index, which in turn affects the reflected wavelengths that produce structural color in butterfly wing scales(63). This implies that the interaction between pigments and CPs can be very subtle, resulting in the formation of unique nanostructures, such as those on wing scales, that produce brilliant structural colors.

Removed in the revised version: This implies that the interaction between pigments and CPs can be very subtle, resulting in the formation of unique nanostructures, such as those on wing scales, that produce brilliant structural colors.

Removed in the revised version: Consequently, to maintain the accuracy of the color pattern, the localization of CPs in the cuticle may be very important. In a previous study employing microarray analysis, different CPs were found in differently colored areas of the epidermis in Papilio xuthus larvae(64). We investigated the CP genes highly expressed in the black region of Papilio xuthus caterpillars. Thirteen genes had orthologous genes in silkworms (Table S5). Among them, 11 genes (BmorCPR67, BmorCPR71, BmorCPR76, BmorCPR79, BmorCPR99, BmorCPR107, BmorCPT4, BmorCPH5, BmorCPFL4, BmorCPG27 and BmorCPG4) were significantly upregulated in homozygous (mamo-/mamo-) knockout individuals at some time points from the 3rd day of the 4th instar to the 5th instar; BmorCPG4 was also among the 18 previously detected CP genes, and two genes (BmorCPG38 and BmorCPR129) were not different between homozygous (mamo-/mamo-) and heterozygous (mamo-/+) individuals (Fig. S14).

Earlier version: The 28 CP genes mentioned above have significantly upregulated expression in homozygous (mamo-/mamo-) gene knockout individuals at some stages. These CPs may be involved in the transportation or cross-linking of melanin in the cuticle. However, among them, there were no differences in their expression level in some periods compared with the control, and some genes were significantly downregulated at some time points in melanic individuals (Fig. 7). This suggests that the regulation of CP genes is complex and may involve other transcription factors and feedback effects. CPs are essential components of the insect cuticle and are involved in cuticular microstructure construction (65), body shape development(66), wing morphogenesis(67), and pigmentation(53). CP genes usually account for over 1% of the total genes in an insect genome and can be categorized into several families, including CPR, CPG, CPH, CPAP1, CPAP3, CPT, CPF and CPFL(68). The CPR family is the largest group of CPs, containing a chitin-binding domain called the Rebers and Riddiford motif (R&R)(69). The variation in the R&R consensus sequence allows subdivision into three subfamilies (RR-1, RR-2, and RR-3)(70). Among the 28 CPs, 11 RR-1 genes, 6 RR-2 genes, 4 hypothetical cuticular protein (CPH) genes, 3 glycine-rich cuticular protein (CPG) genes, 3 cuticular protein Tweedle motif (CPT) genes, and 1 CPFL (like the CPFs in a conserved C-terminal region) gene were identified. The RR-1 consensus among species is usually more variable than RR-2, which suggests that RR-1 may have a species-specific function. RR-2 often clustered into several branches, which may be due to gene duplication events in co-orthologous groups and may result in conserved functions between species (71). The classification of CPH is due to their lack of known motifs. In the epidermis of Lepidoptera, the CPH genes often have high expression levels. For example, BmorCPH24 had a highest expression level, in silkworm larvae epidermis(72). The CPG protein is rich in glycine. The CPH and CPG genes are less commonly found in insects outside the order Lepidoptera (73). This suggests that they may provide species specific functions for the Lepidoptera. CPT contains a Tweedle motif, and the TweedleD1 mutation has a dramatic effect on body shape in D. melanogaster(74). The CPFL members are relatively conserved in species and may be involved in the synthesis of larval cuticles(75). CPT and CPFL may have relatively conserved functions among insects. The CP genes are a group of rapidly evolving genes, and their copy numbers may undergo significant changes in different species. In addition, RNAi experiments on 135 CP genes in brown planthopper (Nilaparvata lugens) showed that deficiency of 32 CP genes leads to significant defective phenotypes, such as lethal, developmental retardation, etc. It is suggested that the 32 CP genes are indispensable, and other CP genes may have redundant and complementary functions (76). In previous studies, it was found that the construction of the larval cuticle of silkworms requires the precise expression of over two hundred CP genes(22). The production, interaction, and deposition of CPs and pigments are complex and precise processes, and our research shows that Bm-mamo plays an important regulatory role in this process in silkworm caterpillars. For further understanding of the role of CPs, future work should aim to identify the function of important cuticular protein genes and the deposition mechanism in the cuticle.

Revised version: Consequently, to maintain the accuracy of the color pattern, the localization of CPs in the cuticle may be very important. In a previous study employing microarray analysis, different CPs were found in differently colored areas of the epidermis in Papilio xuthus larvae(64). We investigated whether the CP genes were highly expressed in the black region of Papilio xuthus caterpillars. Thirteen orthologous genes were found in silkworms (Table S7). Among them, 11 genes (BmorCPR67, BmorCPR71, BmorCPR76, BmorCPR79, BmorCPR99, BmorCPR107, BmorCPT4, BmorCPH5, BmorCPFL4, BmorCPG27 and BmorCPG4) were significantly upregulated in homozygous (mamo-/mamo-) knockout individuals at some time points from the 3rd day of the 4th instar to the 5th instar; BmorCPG4 was also among the 18 previously detected CP genes, and two genes (BmorCPG38 and BmorCPR129) were not differentially expressed between homozygous (mamo-/mamo-) and heterozygous (mamo-/+) individuals (Fig. S14). The expression of the 28 CP genes mentioned above was significantly upregulated in homozygous (mamo-/mamo-) gene knockout individuals at several stages. These CPs may be involved in the transportation or cross-linking of melanin in the cuticle. However, there were no differences in the expression levels of these genes during some periods compared with those in the control group, and the expression of some genes was significantly downregulated at some time points in the melanic individuals (Fig. 7). This finding suggested that the regulation of CP genes is complex and may involve other transcription factors and feedback effects. CPs are essential components of the insect cuticle and are involved in cuticular microstructure construction (65), body shape development(66), wing morphogenesis(67), and pigmentation(53). CP genes usually account for more than 1% of the total genes in an insect genome and can be categorized into several families, including CPR, CPG, CPH, CPAP1, CPAP3, CPT, CPF and CPFL(68). The CPR family is the largest group of CPs and contains a chitin-binding domain called the Rebers and Riddiford motif (R&R)(69). The variation in the R&R consensus sequence can be divided into three subfamilies (RR-1, RR-2, and RR-3)(70). Among the 28 CPs, 11 RR-1 genes, 6 RR-2 genes, 4 hypothetical cuticular protein (CPH) genes, 3 glycine-rich cuticular protein (CPG) genes, 3 cuticular protein Tweedle motif (CPT) genes, and 1 CPFL (like the CPFs in a conserved C-terminal region) gene were identified. The RR-1 consensus among species is usually more variable than the RR-2 consensus, which suggests that RR-1 may have a species-specific function. RR-2 often clustered into several branches, which may be due to gene duplication events in co-orthologous groups and may result in conserved functions between species (71). The classification of CPH was based on the lack of known motifs. In the epidermis of Lepidoptera, CPH genes often exhibit high expression levels. For example, BmorCPH24 has the highest expression level in the silkworm larval epidermis(72). The CPG protein is rich in glycine. The CPH and CPG genes are less commonly found in insects outside the order Lepidoptera (73). This finding suggested that these genes may provide species-specific functions for Lepidoptera. CPT contains a Tweedle motif, and the TweedleD1 mutation has a dramatic effect on body shape in D. melanogaster(74). The CPFL family members are relatively conserved among species and may be involved in the synthesis of larval cuticles(75). CPT and CPFL may have relatively conserved functions among insects. The CP genes are a group of rapidly evolving genes, and their copy numbers may undergo significant changes in different species. In addition, RNAi experiments on 135 CP genes in the brown planthopper (Nilaparvata lugens) revealed that deficiency of 32 CP genes leads to significant defects in phenotypes, such as lethality and developmental retardation. These findings suggested that the 32 CP genes are indispensable and that other CP genes may have redundant and complementary functions (76). In previous studies, it was found that the construction of the larval cuticle of silkworms requires the precise expression of more than two hundred CP genes(22). The production, interaction, and deposition of CPs and pigments are complex and precise processes, and our research showed that Bm-mamo plays an important regulatory role in this process in silkworm caterpillars. To further understand the role of CPs, future work should aim to identify the functions of important cuticular protein genes and the deposition mechanism in the cuticle.

Heatmap of 28 differentially expressed cuticular protein genes.

Earlier version: Differentially expressed cuticular protein genes between homozygotes (mamo-/mamo-) and heterozygotes (mamo-/+) of Bm-mamo knockout line individuals. Red indicates upregulation in homozygous individuals. Blue indicates downregulation in homozygous individuals. White indicates no difference in expression level or no investigation. 4th1d indicates 4-instar day one, M indicates molting, and h indicates hours.

Revised version: Differentially expressed cuticular protein genes between homozygous (mamo-/mamo-) and heterozygous (mamo-/+) Bm-mamo knockout individuals. Red indicates upregulation in homozygous individuals. Blue indicates downregulation in homozygous individuals. White indicates no difference in expression level or no investigation. 4th1d indicates 4-instar Day 1, M indicates molting, and h indicates hours.

Added in the revised version: Maturation of pigment granules

Earlier version: In addition, among the 191 DEGs found in the comparative transcriptome data, we also discovered some interesting genes. For example, BGIBMGA013242, encoding a major facilitator superfamily protein (MFS), named BmMFS, which is responsible for the cheek and tail spot (cts) mutant, was significantly upregulated in the bd mutant. Deficiency of this gene (BmMFS) results in chocolate-colored head and anal plates of the silkworm caterpillar(77). In the ommochrome metabolic pathway, Bm-re, encoding an MFS protein, may function in the transportation of some amino acids, such as cysteine or methionine, into pigment granules(78). Therefore, the encoded product of BmMFS may participate in pigmentation by promoting the maturation of pigment granules. Moreover, BGIBMGA013576 and BGIBMGA013656, which encode MFS domain-containing proteins belonging to the solute carrier family 22 (SLC22) and solute carrier family 2 (SLC2) families, respectively, were significantly upregulated in bd. These two genes may participate in pigmentation in a similar manner to the BmMFS gene. In addition, the red Malpighian tubules (red) gene encodes LysM domain-containing proteins. Deficiency of red results in a significant decrease in orange pterin in the wing of Colias butterflies. Research suggests that the product of red may interact with V-ATPase to modulate vacuolar pH across a variety of endosomal organelles, thereby affecting the maturation of pigment granules in cells(79). This indicates that the maturation of pigment granules plays an important role in the coloring of Lepidoptera, and Bm-mamo may be involved in regulating the maturation of pigment granules in the epidermal cells of silkworms.

Revised version: In addition, among the 191 DEGs found in the comparative transcriptome data, we also discovered several interesting genes. For example, BGIBMGA013242, which encodes a major facilitator superfamily protein (MFS) named BmMFS and is responsible for the cheek and tail spot (cts) mutant, was significantly upregulated in the bd mutant. Deficiency of this gene (BmMFS) results in chocolate-colored head and anal plates on the silkworm caterpillar(77). In the ommochrome metabolic pathway, Bm-re, which encodes an MFS protein, may function in the transportation of several amino acids, such as cysteine or methionine, into pigment granules(78, 79). Therefore, the encoded product of BmMFS may participate in pigmentation by promoting the maturation of pigment granules. Moreover, BGIBMGA013576 and BGIBMGA013656, which encode MFS domain-containing proteins belonging to the solute carrier family 22 (SLC22) and solute carrier family 2 (SLC2) families, respectively, were significantly upregulated in bd. These two genes may participate in pigmentation in a similar manner to that of the BmMFS gene. In addition, the red Malpighian tubule (red) gene encodes a LysM domain-containing protein. Red deficiency results in a significant decrease in orange pterin in the wing of Colias butterflies. Research suggests that the product of red may interact with V-ATPase to modulate vacuolar pH across a variety of endosomal organelles, thereby affecting the maturation of pigment granules in cells(80). These findings indicate that the maturation of pigment granules plays an important role in the coloring of Lepidoptera and that Bm-mamo may be involved in regulating the maturation of pigment granules in the epidermal cells of silkworms.

Earlier version: Bm-mamo may affect the synthesis of melanin in epidermis cells by regulating yellow, DDC, and tan; regulate the maturation of melanin granules in epidermis cells through BmMFS; and affect the deposition of melanin granules in the cuticle by regulating CP genes, thereby comprehensively regulating the color pattern in caterpillars.

Revised version: Bm-mamo may affect the synthesis of melanin in epidermal cells by regulating yellow, DDC, and tan, regulate the maturation of melanin granules in epidermal cells through BmMFS, and affect the deposition of melanin granules in the cuticle by regulating CP genes, thereby comprehensively regulating the color pattern of caterpillars.

Added in the revised version: The upstream sequence of Bm-mamo

Earlier version: Moreover, in D. melanogaster, mamo is needed for functional gamete production. In the silkworm, Bm-mamo has a conserved function in female reproduction. However, mamo has developed new functions in color pattern regulation in silkworm caterpillars. We found that in D. melanogaster, the mamo gene is mainly expressed during the adult stage; it is not expressed during the 1st instar or 2nd instar larval stage and has very low expression at the 3rd instar (Table S6). In addition, several binding sites of mamo were found near the TTS of yellow in D. melanogaster (Fig. S15). The yellow gene is a key melanin metabolic gene with upstream and intron sequences that have been identified as multiple CREs, and it has been considered a research model for CREs (80, 81). This may be due to a change in the expression pattern of mamo, causing this gene to develop a new function of regulating coloration in silkworm caterpillars.

Revised version: Moreover, in D. melanogaster, mamo is needed for functional gamete production. In the silkworm, Bm-mamo has a conserved function in female reproduction. The DNA-binding motifs of the mamo proteins in silkworms and fruit flies are highly conserved, suggesting that they may regulate the same downstream target genes. However, mamo has developed new functions in color pattern regulation in silkworm caterpillars. We found that in D. melanogaster, the mamo gene is expressed mainly during the adult stage; it is not expressed during the 1st instar or 2nd instar larval stage and has very low expression at the 3rd instar (Table S8). In addition, several binding sites of mamo were found near the TSS of yellow in D. melanogaster (Fig. S15). The yellow gene is a key melanin metabolic gene with upstream and intron sequences that have been identified as multiple CREs, and it has been considered a research model for CREs (81, 82). This difference may be due to a change in the expression pattern of mamo, which led to the development of a new function for this gene in regulating coloration in silkworm caterpillars. Changes in gene expression patterns are generally believed to be the result of the evolution of CREs. Because CREs play important roles in the spatiotemporal expression pattern regulation of genes, many CREs are found in noncoding regions, which are relatively prone to sequence variation compared with the sequences of coding genes(49). However, the molecular mechanism underlying the sequence changes in CREs is unclear.

Added in the revised version: TFs, because they recognize relatively short sequences, generally between 4 and 20 bp in length, can have many binding sites in the genome(83). Therefore, one member of the TF family has the potential to regulate many genes in one genome. Single regulation of specific TFs can lead to patterned changes in the expression of multiple downstream genes, enabling organisms to adapt to the environment by altering a given type of trait. For example, the marine form of the three-spined stickleback (Gasterosteus aculeatus) has thick armor, whereas the lake population (which was recently derived from the marine form) does not. Research has shown that pelvic loss in different natural populations of three-spined stickleback fishes occurs via regulatory mutations resulting from the deletion of a tissue-specific enhancer (Pel) of the pituitary homeobox transcription factor 1 (Pitx1) gene. The researchers genotyped 13 pelvic-reduced populations of three-spined sticklebacks from disparate geographic locations. Nine of the 13 pelvic-reduced stickleback populations had sequence deletions of varying lengths, all of which were located at the Pel enhancer (84).

Added in the revised version: We investigated nucleotide diversity in 51 wild silkworms and 171 domesticated silkworms (34, 85). The nucleotide diversity of introns and upstream sequences of Bm-mamo in wild silkworms was significantly greater than that in domestic silkworms. In addition, the approximately 1 kb genomic region upstream had high fixation indices (FST) (Fig. S16). This indicates a significant degree of differentiation between the wild strains in this genomic region. Multiple sequence alignment analysis of this region was performed between 12 wild silkworms and 12 domestic silkworms (Fig. S17 & Table S9). The results showed that the wild silkworm group had a greater degree of nucleotide variation, while the sequences in domestic silkworms were highly conserved. The domestic silkworm group had two different forms: one had a long interspersed nuclear element (LINE) inserted at approximately 4.5 kb, and the other did not have this transposon sequence. It is suggested that this genomic region is prone to variation in wild silkworm and is fixed in domestic silkworms. The larvae of domesticated strains often have a lighter or even pale body color, but the wild-type strains have a darker color. The Bm-mamo gene may be involved in the domestication of silkworms.

Added in the revised version: The function of TFs is to precisely regulate the temporal and spatial expression of target genes. Therefore, the expression of TFs requires strict regulation. There are long intergenic regions upstream of many important TFs, dozens of kilobase pairs (Kb) to hundreds of Kb, which may contain many CREs for better control its expression pattern. It has often been believed that changes in CREs are caused by random mutations. However, recent studies have shown that there is a mutation bias in the genome; compared with that in the intergenic region, the mutation frequency is reduced by half inside gene bodies and by two-thirds in essential genes. In addition, they compared the mutation rates of genes with different functions. The mutation rate in the coding region of essential genes (such as translation) is the lowest, and the mutation rates in the coding region of specialized functional genes (such as environmental response) are the highest. These patterns are mainly affected by the features of the epigenome(86). Due to the plasticity of epigenomic features, mutations bias associated with epigenomes may even lead to environmental influences on mutations (87). This finding suggested that different functional regions in the genome are subject to distinct controls and that some sequences can undergo procedural changes under environmental changes. Therefore, some variations in the CRE of TFs in response to environmental changes may be mutation bias.

Removed in the revised version: It is generally believed that changes in gene expression patterns are the result of the evolution of CREs. Because CREs have important functions in the spatiotemporal expression pattern regulation of genes, many are found in the noncoding region, which is relatively prone to sequence variation compared with the sequences of coding genes(49). However, the molecular mechanism underlying the sequence change of CREs is not clear.

Removed in the revised version: Transcription factors (TFs), because they recognize relatively short sequences, generally between 4 and 20 bp in length, can have many binding sites in the genome(82). Therefore, one member of the TF family has the potential to regulate almost all genes in one genome. For example, in the PWM sequence of the mamo protein recognition sequence, 10,622 genes (∼73% of the total number of genes) contain binding sites within 2 kb of their up/downstream region in silkworms (Table S7 and Table S8). Divergence of CREs is currently considered an important cause of phenotypic evolution(83). For example, the marine form of the three-spined stickleback (Gasterosteus aculeatus) has thick armor, and the lake population (which was recently derived from the marine form) does not. Research has shown that pelvic loss in different natural populations of three-spined stickleback fish occurs by regulatory mutations deleting a tissue-specific enhancer (Pel) of the pituitary homeobox transcription factor 1 (Pitx1) gene. The researchers genotyped 13 pelvic-reduced populations of three-spined stickleback from disparate geographic locations. Nine of the 13 pelvic-reduced stickleback populations had sequence deletions of varying lengths, all of which were located at the Pel enhancer. The author suggested that the Pitx1 locus of the stickleback genome may be prone to double-stranded DNA breaks that are subsequently repaired by nonhomologous end joining (NHEJ) (84). There are also examples of this in Lepidoptera. The cortex gene encodes a member of the Cdc20/fizzy family of cell cycle regulators. The cortex gene locus is frequently mutated in Lepidoptera, such as the black (carbonaria) form of the peppered moth (Biston betularia) in industrial melanism(85), the diversification of the leaf wing of oakleaf butterflies(86), the mimicry wing patterning in Heliconius butterflies(87), and the plastic wing pattern in Junonia coenia (88). The explanation of DNA fragility is very interesting, but it still cannot explain why DNA sequence mutations occur frequently in a specially site corresponding to the adaptive phenotype. DNA fragility sites are widely distributed in the genome(89). Under the stimulation of the same environmental conditions, all fragile sites are thought to undergo frequent mutations, and this clearly brings adverse burdens to the organism.

Removed in the revised version: Some research suggests that common fragile sites (CFSs) are specific regions in the genome of all individuals that are prone to DNA double stranded breaks (DSBs) and subsequent rearrangements, and the DSB sites are correlated with epigenetic markers for chromatin accessibility, including DNaseI HSSs, H3K4me3, and CTCF binding sites(90, 91).

Removed in the revised version: Moreover, some studies have shown that there is a mutation bias in the genome; compared with the intergenic region, the mutation frequency is reduced by half inside gene bodies and by two-thirds in essential genes. In addition, they compared the mutation rates of genes with different functions. The mutation rate in the coding region of essential genes (such as translation) is the lowest, and the mutation rates in the coding region of specialized functional genes (such as environmental response) are the highest. These patterns are mainly affected by the features of the epigenome(92). Epigenetic modifications include DNA methylation, histone modifications, remodeling of nucleosomes and higher order chromatin reorganization(93). TFs, such as BTB-ZF proteins, can recruit histone modification factors such as DNA methyltransferase 1 (DMNT1), cullin3 (CUL3), histone deacetylase 1 (HDAC1), and histone acetyltransferase 1 (HAT1) to perform chromatin remodeling at specific genomic sites based on specific DNA recognition capabilities(94). In addition, programmed events related to DSBs and repair at genome-specific sites have been reported. For example, PRDM9 can determine recombination hotspots by H3K4 and H3K36 trimethylation (H3K4me3 and H3K36me3) of nucleosomes near its DNA-binding site in meiosis recombination(95). This suggests that, on the basis of randomness, TFs and epigenomic features can determine the mutation bias at unique positions of the genome.

Earlier version: Materials and Methods Silkworm strains

Revised version: Materials and Methods

Added in the revised version: Silkworm strains

Earlier version: The bd and bdf mutant strains and wild-type strains Dazao and N4 were obtained from the bank of genetic resources of Southwest University. Silkworms were reared on mulberry leaves or artificial diets at 25 °C and 73% relative humidity under dark for 12 hours and light for 12 hours.

Revised version: The bd and bdf mutant strains and the wild-type Dazao and N4 strains were obtained from the bank of genetic resources of Southwest University. Silkworms were reared on mulberry leaves or artificial diets at 25°C and 73% relative humidity in the dark for 12 hours and light for 12 hours.

Positional cloning of the bd locus

Earlier version: For mapping of the bd locus, F1 heterozygous individuals were obtained from a cross between a bdf strain and a Dazao strain. Then, an F1 female was crossed with a bdf male (BC1F), and an F1 male was backcrossed with a bdf female (BC1M). A total of 1162 BC1M individuals were used for recombination analysis. Genomic DNA was extracted from the parents (Dazao, bdf and F1) and each BC1 individual using the phenol chloroform extraction method. Available DNA molecular markers were sought through polymorphism screening and linkage analysis. The primers used for mapping are listed in Table S9.

Revised version: For mapping of the bd locus, F1 heterozygous individuals were obtained from a cross between a bdf strain and a Dazao strain. Then, an F1 female was crossed with a bdf male (BC1F), and an F1 male was backcrossed with a bdf female (BC1M). A total of 1162 BC1M individuals were used for recombination analysis. Genomic DNA was extracted from the parents (Dazao, bdf and F1) and each BC1 individual using the phenol chloroform extraction method. Available DNA molecular markers were identified through polymorphism screening and linkage analysis. The primers used for mapping are listed in Table S10.

Phylogenetic analysis

Earlier version: To determine whether Bm-mamo orthologs existed in other species, the BlastX program of the National Center for Biotechnology Information (NCBI) (http://www.ncbi.nlm.nih.gov/BLAST/) was used. The sequences of Bm-mamo-L and Bm-mamo-S were blasted against the nonredundant protein sequence (nr) database. Sequences with a maximum score and E-value≤10-4 were downloaded. The sequences of multiple species were subjected to multiple sequence alignment of the predicted amino acid sequences by MUSCLE(96). Then, a phylogenetic tree was constructed using the neighbor-joining method with the MEGA7 program (Pearson model). The confidence levels for various phylogenetic lineages were estimated by bootstrap analysis (2,000 replicates).

Revised version: To determine whether Bm-mamo orthologs were present in other species, the BlastX program of the National Center for Biotechnology Information (NCBI) (http://www.ncbi.nlm.nih.gov/BLAST/) was used. The Bm-mamo-L and Bm-mamo-S sequences were subjected to BLAST searches against the nonredundant protein sequence (nr) database.

Added in the revised version: Sequences with a maximum score and an E-value≤10-4 were downloaded. The sequences of multiple species were subjected to multiple sequence alignment of the predicted amino acid sequences by MUSCLE(88). A phylogenetic tree was subsequently constructed using the neighborLjoining method with the MEGA7 program (Pearson model). The confidence levels for various phylogenetic lineages were estimated by bootstrap analysis (2,000 replicates).

Quantitative PCR

Earlier version: Total RNA was isolated from the whole body and integument of the silkworms using TRIzol reagent (Invitrogen, California, USA), purified by phenol chloroform extraction and then reverse transcribed with a PrimeScript™ RT Reagent Kit (TAKARA, Dalian, China) according to the manufacturers’ protocol. qPCR was performed using a CFX96™ Real-Time PCR Detection System (Bio-Rad, Hercules, CA) with a Real-Time PCR System and an iTaq Universal SYBR Green Supermix System (Bio-Rad). The cycling parameters were as follows: 95 °C for 3 min, followed by 40 cycles of 95 °C for 10 s and annealing for 30 s. The primers used for target genes are listed in Table S9.

Revised version: Total RNA was isolated from the whole body and integument of the silkworms using TRIzol reagent (Invitrogen, California, USA), purified by phenol chloroform extraction and then reverse transcribed with a PrimeScript™ RT Reagent Kit (TAKARA, Dalian, China) according to the manufacturers’ protocol.

Earlier version: Expression levels for each sample were determined with biological replicates, and each sample was analyzed in triplicate. The gene expression levels were normalized against the expression levels of the ribosomal protein L3 (RpL3). The relative expression levels were analyzed using the classical R=2−ΔΔCt method.

Revised version: qPCR was performed using a CFX96™ Real-Time PCR Detection System (Bio-Rad, Hercules, CA) with a qPCR system and an iTaq Universal SYBR Green Supermix System (Bio-Rad). The cycling parameters were as follows: 95°C for 3 min, followed by 40 cycles of 95°C for 10 s and annealing for 30 s. The primers used for the target genes are listed in Table S10. The expression levels of the genes in each sample were determined with biological replicates, and each sample was analyzed in triplicate. The gene expression levels were normalized against the expression levels of the ribosomal protein L3 (RpL3). The relative expression levels were analyzed using the classical R=2−ΔΔCt method.

siRNA for gene knockdown

Earlier version: siRNAs for Bm-mamo were designed by the siDirect program (http://sidirect2.rnai.jp). The target siRNAs and negative control were synthetized by Tsingke Biotechnology Company Limited. The siRNA (5 μl, 1 μl/μg) was injected from the abdominal spiracle into the hemolymph at the fourth instar (day 3) larval stage. Immediately after injection, phosphate-buffered saline (pH 7.3) droplets were placed nearby, and a 20-voltage pulse for one second and pause for one second were repeated 3 times. The phenotype was observed for fifth-instar larvae. The left and right epidermis were separately dissected from the injected larvae, and RNA was extracted. Then, cDNA was synthetized, and the expression level of the gene was detected by qPCR.

Revised version: siRNAs for Bm-mamo were designed with the siDirect program (http://sidirect2.rnai.jp). The target siRNAs and negative controls were synthesized by Tsingke Biotechnology Company Limited. The siRNA (5 μl, 1 μl/μg) was injected from the abdominal spiracle into the hemolymph at the fourth instar (Day 3) larval stage. Immediately after injection, phosphate-buffered saline (pH 7.3) droplets were placed nearby, and a 20-voltage pulse for one second and pause for one second were applied 3 times. The phenotype was observed for fifth-instar larvae. The left and right epidermis were separately dissected from the injected larvae, after which RNA was extracted. Then, cDNA was synthesized, and the expression level of the gene was detected via qPCR.

sgRNA synthesis and RNP complex assembly

Earlier version: CRISPRdirect (http://crispr.dbcls.jp/doc/) online software was used to screen appropriate single guide RNA (sgRNA) target sequences. The gRNAs were synthetized by Beijing Genomics Institute. sgRNA templates were transcribed using T7 polymerase with RiboMAX™ Large-Scale RNA Production Systems (Promega, Beijing, China) according to the manufacturer’s instructions. The RNA transcripts were purified using 3 M sodium acetate (pH 5.2) and anhydrous ethanol (1:30) precipitation, washed using 75% ethanol, and eluted in nuclease-free water. All injection mixes contained 300 ng/µL Cas9 nuclease (Invitrogen, California, USA) and 300 ng/µL purified sgRNA. Before injection, mixtures of Cas9 nuclease and gRNA were incubated for 15 min at 37 °C to reconstitute active ribonucleoproteins (RNPs).

Revised version: CRISPRdirect (http://crispr.dbcls.jp/doc/) online software was used to screen appropriate single guide RNA (sgRNA) target sequences. The gRNAs were synthetized by the Beijing Genomics Institute. The sgRNA templates were transcribed using T7 polymerase with RiboMAX™ Large-Scale RNA Production Systems (Promega, Beijing, China) according to the manufacturer’s instructions. The RNA transcripts were purified using 3 M sodium acetate (pH 5.2) and anhydrous ethanol (1:30) precipitation, washed with 75% ethanol, and eluted in nuclease-free water. All injection mixes contained 300 ng/µL Cas9 nuclease (Invitrogen, California, USA) and 300 ng/µL purified sgRNA. Before injection, mixtures of Cas9 nuclease and gRNA were incubated for 15 min at 37°C to reconstitute active ribonucleoproteins (RNPs)(89).

Microinjection of embryos

Earlier version: For embryo microinjection, microcapillary needles were manufactured using a PC-10 model micropipette puller (Narishige, Tokyo, Japan). Microinjection was performed using Eppendorf’s TransferMan NK 2 and a FemtoJet 4i system (Eppendorf, Hamburg, Germany). The eggs used for microinjection came from the mating of female and male wild-type Dazao moths. Within 4 hours of being laid, the eggs were adhered onto a clean glass slide. CRISPR/Cas9-messenger RNP mixtures with volumes of approximately 1 nL were injected into the middle of the eggs, and the wound was sealed with glue. All injected embryos were allowed to develop in a climate chamber at 25 °C and 80% humidity.

Revised version: For embryo microinjection, microcapillary needles were manufactured using a PC-10 model micropipette puller (Narishige, Tokyo, Japan). Microinjection was performed using an Eppendorf TransferMan NK 2 and a FemtoJet 4i system (Eppendorf, Hamburg, Germany). The eggs used for microinjection were generated by mating female and male wild-type Dazao moths. Within 4 hours of culture, the eggs were allowed to adhere to a clean glass slide. CRISPR/Cas9-messenger RNP mixtures with volumes of approximately 1 nL were injected into the middle of the eggs, and the wound was sealed with glue. All the injected embryos were allowed to develop in a climate chamber at 25°C and 80% humidity(90).

Comparative genomics

The reference genome of Dazao was downloaded from Silkbase (https://silkbase.ab.a.u-tokyo.ac.jp/cgi-bin/index.cgi). The bdf genome was obtained from the silkworm pangenome project(34).

Earlier version: The short reads of the bdf strains were mapped to the silkworm reference genome by BWA55 v0.7.17 mem with default parameters. The SAMtools56 v1.11 and Picard v2.23.5 (https://broadinstitute.github.io/picard/) programs were used to filter the unmapped and duplicated reads. A GVCF file of the samples was obtained using GATK57 v4.1.8.1 HaplotypeCaller with the parameter –ERC = GVCF. The VCF files of insertion/deletions (indels) and single-nucleotide polymorphisms (SNPs) were used for further analysis by eGPS software.

Revised version: The short reads of the bdf strains were mapped to the silkworm reference genome by BWA55 v0.7.17 mem with default parameters. The SAMtools56 v1.11 and Picard v2.23.5 (https://broadinstitute.github.io/picard/) programs were used to filter the unmapped and duplicated reads. A GVCF file of the samples was obtained using GATK57 v4.1.8.1 HaplotypeCaller with the parameter-ERC = GVCF. The VCF files of insertions/deletions (indels) and single-nucleotide polymorphisms (SNPs) were used for further analysis via eGPS software.

Downstream target gene screening

Online software (http://zf.princeton.edu) was used to predict the DNA-binding site for Cys2His2 zinc finger proteins. The confident ZF domains with scores higher than 17.7 were chosen. RF regression on the B1H model was used to predict the DNA-binding sites. Then, the sequence logo and position weight matrices (PWMs) of the DNA-binding sites were obtained. The sequences 2 kb upstream and downstream of the predicted genes were extracted by a Perl script. The FIMO package of the MEME suite was used to search for binding sites in the silkworm genome.

Analysis of EMSA

Earlier version: Primers with binding sites and their flanking sequences were designed according to the FIMO results. A biotin label was added to one end of the upstream primer (probe), and the downstream primer was normal. Primers with the same sequence as the labeled probes were used as competitive probes. The EMSA experiment was conducted with an EMSA reagent kit (Beyotime, Shanghai, China) according to the manufacturers’ protocol.

Revised version: Primers with binding sites and flanking sequences were designed according to the FIMO results. A biotin label was added to one end of the upstream primer (probe), and the downstream primer was used as a reference. Primers with the same sequence as the labeled probes were used as competitive probes. The EMSA was conducted with an EMSA reagent kit (Beyotime, Shanghai, China) according to the manufacturers’ protocol.

Removed in the revised version: Parthenogenesis

Removed in the revised version: Virgin moths were dissected in 0.9% bacterium-free physiological saline. The ovaries were extracted, and water was removed by clean filter paper. The ovaries were stored at 25 °C and 40% relative humidity for 10 hours. The eggs were treated in 46 °C warm water for 18 minutes and then transferred to room temperature for 5-10 minutes. The activated egg was protected in an environment of 15-17 °C and 18% relative humidity for 3 days. Then, it was immersed in acid immediately (with a hydrochloric acid specific gravity of 1.075 at 46 °C for 5 minutes) or transferred to the same conditions as the hibernating eggs for protection. The average hatching rate of silkworm eggs treated in this way was 50-60%.

Supporting information